Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# El ordenador juega al Wordle

En este segundo apartado vamos a hacer un programa que juege al wordle. El ordenador va a tratar de encontrar la palabra que nosotros hayamos escogido. Para no hacer trampas debemos escoger una palabra que se encuentre el el fichero <code>lemario.txt</code>. Vamos a coger una palabra de 5 letras, si quisieramos cambiar esa longitud deberíamos cambiar las constantes de forma adecuada. Por ahora vamos a coger las constantes que ya teníamos definidas.

```python
LETTER_IN_POS = 0
LETTER_IN_NO_POS = 1
LETTER_NOT_IN = 2
MAX_ATTEMPTS = 6
WORD_LENTGH = 5
```

De entrada, tendremos que usar las funciones que ya teníamos en el apartado anterior.

<!-- grade_id: cell-988ffeb76b0056bd, points: 0, grade: false, solution: true -->
```python
LETTER_IN_POS = 0
LETTER_IN_NO_POS = 1
LETTER_NOT_IN = 2
MAX_ATTEMPTS = 6
WORD_LENTGH = 5

def read_words(filename, length):
    """
    This function reads the filename. Each line contains a word.
    This function returns a list with the words with the indicated length 
    """
    # YOUR CODE HERE
    raise NotImplementedError()


def get_clues(secret, user_word):
    """ 
    This funcction returns the clues for the word user_word with respect secret. 
    For each letter it returns 
      - LETTER_IN_POS if the letter is in secret in the same position. 
      - LETTER_IN_NO_POS if the letter is in secret but in another position.
      - LETTER_NOT_IN if the letter is not in secret.
    The function returns a list with the corresponding values.
    Example
    secret = 'GOTAS'
    user_word = 'GALGO'
    returns -> [LETTER_IN_POS, LETTER_IN_POS, LETTER_NOT_IN, LETTER_IN_NO_POS, LETTER_IN_NO_POS]
    """
    # YOUR CODE HERE
    raise NotImplementedError()

def are_all_in_pos(answer):
    """
    This function returns a boolean.
    It returns True if all elements in the list answer are LETTER_IN_POS.
    """
    # YOUR CODE HERE
    raise NotImplementedError()


```

## El algoritmo

Veamos someramente el algoritmo que va a seguir nuestro programa. Tenemos todas las posibles palabras almacenadas en una lista. El programa escoge la primera palabra de la lista y la propone como posible candidata. Pregunta por la respuesta y la almacena. A continuación, siguiendo por la siguiente palabra de la lista, busca la primera palabra que es <em>compatible</em> con la respuesta. La propone y la almacena en la lista la palabra y la respuesta. En todo momento lo que se hace es buscar a partir de la última palabra propuesta, la siguiente que es <em>compatible</em> con las respuestas almacenadas. Pero ¿qué significa exactamente ser compatible? Una palabra es compatible con todas las respuestas, si asumiendo que ella es la palabra secreta, da la misma respuesta a las palabras almacenadas.
El programa llevará una lista con las palabras que han sido propuestas y sus respuestas. Por ejemplo, si tenemos las siguientes palabras con sus respuestas


Palabra | Respuesta
--------|------------------
'ababa' |  [0, 2, 1, 2, 0]
'acema' |  [0, 2, 0, 2, 0]

La palabra 'acuma' no es compatible por que la resupuesta que daría a 'acema' es <code>[0,0,2,0,0]</code> en lucar de <code>[0,2,0,2,0]</code>. La palabra 'aleda' sí que es compatible. 



# El programa principal

Para desarrollar el programa empezaremos otra vez por la función principal del programa. Será parecida a la del programa anterior. Necesitamos el almacén de palabras y respuestas para poder averigüar la siguiente palabra compatible, a esa variable la llamaremos <code>storage</code>. Luego necesitamos la lista de palabras y la posición actual para poder seguir las palabras en orden.                           

En cada iteración buscamos la siguiente palabra compatible y almacenamos la nueva respuesta en el almacén. Por último un pequeño detalle, en un juego <em>perfecto</em>, el fácil demostar que el algoritmo anterior llega a la solución. El problema es que el da las respuestas y bien se puede equivocar o directamente intentar hacer trampas. En esos casos, es posible que no exista la siguiente palabra compatible. 


```python
def computer(words_file):
    pos = 0
    words = read_words(words_file, WORD_LENTGH)
    storage = []
    game_cont = True
    while game_cont:
        pos = next_candidate(storage, words, pos)
        if pos<len(words):
            candidate = words[pos]
            answer = get_answer(candidate)
            add_answer(storage, candidate, answer)
            pos += 1
        game_cont = pos<len(words) and not are_all_in_pos(answer)
    if pos==len(words):
        print('You have chosen an illegal word or have given an incorect answer')
    else:
        print(f'The secret word is {candidate}')
    
```

El programa anterior no funciona porque falta la definición de las varias funciones: <code>next_candidate</code>, <code>add_answer</code> y <code>get_answer</code>. Las funciones <code>read_words</code> y <code>are_all_in_pos</code> ya están definidas de la versión anterior. 

En primer lugar daremos la función <code>add_answer</code>. Como parámetros de entrada hay una palabra candidata y debe devolver una lista de 0, 1 y 2 de la misma longitud de la palabra con la respuesta del usuario. El usuario responde con una secuencia de números separada por espacios. 

Se debe comprobar que la respuesta que da el usuario es la correcta, si no lo es se debe preguntar de nuevo al usuario. 

No te olvides de escribir la documentación de la función.

<!-- grade_id: get_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

En la casilla de que sigue comprueba que tu programa funciona correctamente. Escribe tus propias pruebas

```python

```

<!-- grade_id: test_get_answer, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
input_ori = input
class MyInput():
    def __init__(self, answers):
        self.pos = 0
        self.answers = answers
    def input(self, msg):
        if self.pos<len(self.answers):
            ans = self.answers[self.pos]
            self.pos += 1
        else:
            ans = ''
        return ans
try:
    test = ['0 1 0 1 2']
    print(f'checking {test}....', end='')
    myinput = MyInput(test)
    input = myinput.input
    answer = get_answer("abcde")
    assert answer==[0, 1, 0, 1, 2], f'La respuesta es {answer}'
    print('ok')

    test = ['0 0 1 2', '0 1 0 1 2']
    print(f'checking {test}....', end='')
    myinput = MyInput(test)
    input = myinput.input
    answer = get_answer("abcde")
    assert answer==[0, 1, 0, 1, 2], f'La respuesta es {answer}'
    print('ok')


    test = ['0 0 0 3 2', '0 1 0 1 2']
    print(f'checking {test}....', end='')
    myinput = MyInput(test)
    input = myinput.input    
    answer = get_answer("abcde")
    assert answer==[0, 1, 0, 1, 2], f'La respuesta es {answer}'
    print('ok')


    test = ['0 0 0 1 0 2', '0 1 0 1 2']
    print(f'checking {test}....', end='')
    myinput = MyInput(test)
    input = myinput.input    
    answer = get_answer("abcde")
    assert answer==[0, 1, 0, 1, 2], f'La respuesta es {answer}'
    print('ok')


finally:
    input = input_ori
### END HIDDEN TEST
```

A continuación vamos a por la función <code>add_answer</code>. Debe modifificar la variable <code>storage</code> que tiene como parámetro para que incorpore una nueva tupla con la palabra y la respuesta. El primer elemento de la tupla debe ser la palabra y la segunda la respuesta de forma que la variable storage tenga un contenido similar al siguiente

    <code>[('ababa', [0, 2, 1, 2, 0]), ('acema', [0, 2, 0, 2, 0]), ('aleda', [0, 2, 0, 2, 0]) ]</code>

Recuerda de escribir la documentación de la función

<!-- grade_id: add_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

En la casilla a continuación escribe las pruebas que hagas para ver que esa función está bien

<!-- grade_id: cell-bbfd707261b83342, points: 1, grade: true, solution: false -->
```python

```

Para realizar la función <code>next_candidate</code> debebmos partir de la posición actual y encontrar la siguiente palabra compatible con el almacen en la lista de palabras. Por ello vamos a hacer primero una función que nos diga si una palabra es compatible con las almacenadas en el almacén. 

<!-- grade_id: is_compatible, points: 0, grade: false, solution: true -->
```python
def is_compatible(storage, word):
    # YOUR CODE HERE
    raise NotImplementedError()
```

En la siguiente casilla puedes probar tu solución 

```python

```

<!-- grade_id: test_is_compatible, points: 1, grade: true, solution: false -->
```python

```

Ahora puedes dar la implementación de la función <code>next_candidate</code>. Fíjate en el lugar en el que se usa, te indicará los parámetros que tiene y lo que tiene que devolber. No te olvides de poner los comentarios adecuados en la definición

<!-- grade_id: next_candidate, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

En la siguiente casilla debes probar tu solución.

```python

```

<!-- grade_id: test_next_candidate, points: 1, grade: true, solution: false -->
```python

```

## Programa completo

Ha llegado el meomento de hacer el programa completo. Coge la función <code>computer</code> y todas las demás y ponlas a conticunación. La ejecuciön del programa genera un fichero <code>computer.py</code>

<!-- grade_id: sol_completa, points: 0, grade: false, solution: true -->
```python
%%file computer.py

# YOUR CODE HERE
raise NotImplementedError()



```

Ejecuta tu programa en la siguiente casilla para ver si funciona bien

```python

```

<!-- grade_id: test_sol_completa, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TEST
import sys 
import io

stdout = sys.stdout
input_ori = input
class MyInput():
    def __init__(self, answers):
        self.pos = 0
        self.answers = answers
    def input(self, msg):
        if self.pos<len(self.answers):
            ans = self.answers[self.pos]
            self.pos += 1
        else:
            ans = ''
        return ans
try:
    myinput = MyInput(['0 2 1 2 1', 
                       '0 2 2 2 1', 
                       '0 2 2 0 2', 
                       '0 2 2 0 2',
                       '0 2 2 0 0', 
                       '0 1 0 0 0', 
                       '0 0 0 0 0'])
    
    output = io.StringIO("")
    input = myinput.input
    sys.stdout = output
    computer("lemario.txt")
    assert output.getvalue() == 'The secret word is avion\n', f'"{output.getvalue()}"'

    myinput = MyInput(['0 2 1 2 1', 
                       '0 2 2 2 1', 
                       '0 2 2 0 2', 
                       '0 2 2 0 2',
                       '0 2 2 0 0', 
                       '0 1 0 0 0', 
                       '0 0 0 0 1'])
    
    output = io.StringIO("")
    input = myinput.input
    sys.stdout = output
    computer("lemario.txt")
    assert output.getvalue() == 'You have chosen an illegal word or have given an incorect answer\n', f'"{output.getvalue()}"'


    
finally:
    input = input_ori
    sys.stdout = stdout
### END HIDDEN TEST
```

## Ordenador contra sí mismo

Por último haz un programa que juegue contra sí mismo. Haz una función llamada
<code>computer_vs_computer</code> de forma que el ordenador juegue contra sí mismo, sin hacer trampas. La función como las anteriores debe contener un el nombre de un fichero de palabras. Debe devolver una tupla, el primer elemento de la tupla debe ser la palabra seleccionada, el segundo una lista con las todas las respuestas calculadas. Para seleccionar la palabra usa la función <a href='https://docs.python.org/3.8/library/random.html#random.randint'><code>random.randint</code></a>

<!-- grade_id: computer_vs_computer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Ejecuta en la casilla siguientes pruebas para ver que realmente funciona.

```python
computer_vs_computer('lemario.txt')
```

Por último, añade la función <code>computer_cvs_computer</code> al programa <code>computer.py</code>.