Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

[![Creative Commons License](img/CC_80x15.png)](http://creativecommons.org/licenses/by-nc-sa/4.0/) This work is licensed under a [Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License](http://creativecommons.org/licenses/by-nc-sa/4.0/). Author: Simon Pickin 

# Juego de Wordle con clases

En esta práctica vamos a construir una versión basada en el uso de clases del programa para jugar a *Wordle* construido en la práctica previa. Al mismo tiempo, vamos a extender ligeramente su funcionalidad respecto a la práctica anterior.

## Clases principales del  juego

Lo primero que tenemos que hacer es definir la estructura básica de nuestro programa, es decir, decidir cuáles serán las clases principales. Esta decisión es una parte importante del diseño y no es una decisión fácil; de hecho, distintas estructuras son posibles, aunque unas serán mejores que otras. A continuación, tenemos que identificar los métodos de estas clases. Otra vez, esta decisión es una parte importante del diseño de nuestro programa. En esta práctica proporcionamos en el enunciado tanto las clases como las cabeceras de los métodos. A  continuación proporcionamos el esqueleto del código de las clases principales:

<!-- grade_id: cell-ad4c647a0138949c, points: 0, grade: false, solution: true -->
```python
#Lines like the folowing be commented in the final version
#%%writefile wordle_class.py

import platform
if platform.system().lower() == 'windows':
    import os
    os.system("")  # enables ansi escape characters in terminal (supposedly) 
import random

class Game(object):

    # precondition: not (client_is_human==True and server_is_computer==False)
    def __init__(self, filename, client_is_human, server_is_computer, word_length, max_attempts):

        # YOUR CODE HERE
        raise NotImplementedError()

    def play(self):
        
        # YOUR CODE HERE
        raise NotImplementedError()
    
        final_msg = "" if self.server.get_secret() == None \
                    else "The secret word was " + self.server.get_secret() + "."
        GameIO.output(final_msg)
        if not clue.is_winner():
            GameIO.output(f"Game over. Better luck to the {self.player} next time!")
        else:
            GameIO.output(f"Congratulations to the {self.player} for finding the secret word.")

class ComputerServer(object):

    def __init__(self, wordlst):
        # YOUR CODE HERE
        raise NotImplementedError()

    def generate_secret(self):
        self.secret = random.choice(self.wordlst)  # https://docs.python.org/3/library/random.html#random.choice
        
    def get_secret(self):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def get_clue(self, guess):
        # YOUR CODE HERE
        raise NotImplementedError()

    
class HumanServer(object):

    def __init__(self, wordlst):
        # YOUR CODE HERE
        raise NotImplementedError()

    def generate_secret(self):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def get_clue(self, guess):
        # YOUR CODE HERE
        raise NotImplementedError()

    
class HumanPlayer(object):

    # Errors
    UNKNOWN_WORD_MSG = "The word you guessed is not in the dictionary"
    # Warnings
    MSG_PREFIX = "Warning: "
    REPEAT_WORD_MSG = MSG_PREFIX + "you have already guessed that word"
    INCONSISTENT_WORD_MSG = MSG_PREFIX \
            + "the word you guessed is not compatible with the previous clues"

    def __init__(self, wordlst, storage):
        # YOUR CODE HERE
        raise NotImplementedError()

    def next_guess(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def __str__(self):
        # YOUR CODE HERE
        raise NotImplementedError()


class ComputerPlayer(object):

    def __init__(self, wordlst, storage):
        # YOUR CODE HERE
        raise NotImplementedError()

    def next_guess(self):
        # YOUR CODE HERE
        raise NotImplementedError()
        guess = random.choice(compatible_words)
        return guess
        

    def __str__(self):
        # YOUR CODE HERE
        raise NotImplementedError()

```

## Métodos de las clases principales

**Clase Game**
La clase Game es la clase que arranca el juego y gestiona la interacción entre el jugador del juego y el servidor del juego.
- El método `__init__` debe almacenar los argumentos no booleanos en atributos, abrir el fichero y construir la lista de palabras, instanciar el jugador y el servidor según la especificación del segundo y tercer argumentos y, finalmente, llamar a `play`.
- El método `play` contiene el bucle principal del programa; tiene algún parecido con las funciones `human`, `computer` y `computer_vs_computer` de la práctica anterior. Mientras no está terminado el juego (no quedan intentos o se ha adivinado la palabra secreta), se llama a los métodos adecuados del `player` o del `server` (y a un método de otra clase que veremos más adelante) para seguir jugando.

**Clase ComputerServer**
La clase `ComputerServer` representa el papel de servidor de juego desempeñado por el ordenador.

- El método `__init__` es relativamente trivial.
- El método `generate_secret` elige aleatoriamente de la lista de palabras una palabra del tamaño usado en el juego en curso y la asigna al atributo `secret`.
- El método `get_clue` calcula las pistas asociadas a la palabra pasada como argumento. A continuación, imprime las pistas imprime las pistas (usando el médodo `GameIO.clue_output`) y también las devuelve.  Esta función debe hacer uso de los métodos de otras clases como veremos más adelante. En concreto, no debe llamar directamente ni a `print` ni a `input` y debe usar la clase `Clue` que está más adelante.

**Clase HumanServer**
La clase `HumanServer` representa el papel de servidor de juego desempeñado por el usuario.

- El método `__init__` es relativamente trivial.
- El método `generate_secret`. Puesto que con servidor humano, el usuario custodia la palabra secreta solo debe imprimir un mensaje como el siguiente

        Press enter when you've chosen a X-letter word from the word base; but don't tell me what it is!
        
  donde `X` es el número de letras del juego. Para ello debe usar el método adecuado de la clase `GameIO`.
  
- El método `get_clue` pide al usuario que calcule las pistas asociadas a la palabra pasada como parámetro e introduzca las pistas por el teclado. El usuario debería poder introducir una palabra compuesta de cualquier combinación de mayúsculas y minúsculas. A continuación, imprime las pistas (usando el médodo `GameIO.clue_output`) y también las devuelve.  Esta función debe hacer uso de los métodos de otras clases como veremos más adelante. En concreto, no debe llamar directamente ni a `print` ni a `input` y debe usar la clase `Clue` que está más adelante.

**Clase ComputerPlayer**
La clase `ComputerPlayer` representa el papel de jugador de juego desempeñado por el ordenador.

- El método `__init__` es relativamente trivial.
- El método `next_guess` devuelve una palabra elegida aleatoriamente de entre aquellas palabras de la lista de palabras que satisfagan las pistas dadas en pasos anteriores. Levanta una excepción si llega al final del diccionario sin encontrar una palabra adecuada.
- El método `__str__` devuelve la cadena 'computer'. Se llama implícitamente desde los mensajes que se imprimen en pantalla al final de la ejecución (ver el método `play` de la clase `Game`)

**Clase HumanPlayer**
La clase `HumanPlayer` representa el papel de jugador de juego desempeñado por el usuario.

- El método `__init__` es relativamente trivial.
- El método `next_guess` pide al usuario que introduzca una nueva palabra elegida de la lista de palabras. El usuario debería poder introducir una palabra compuesta de cualquier combinación de mayúsculas y minúsculas. Si la palabra introducida no tiene el tamaño adecuado, se imprime el texto como el siguiente:

      The secret word has X letters
   
  Donde `X` es el número de letras del juego. Si la palabra introducida no está en la lista de palabras, se imprime el texto contenido en la variable `UNKNOWN_WORD_MSG`.

  Para ayudar al usuario, también imprimimos una avertencia (sin que se pare la ejecución) en dos casos:
  
    * el jugador humano propone una palabra que ha propuesto previamente en este mismo partido, en cuyo caso, el programa imprime el siguiente mensaje en pantalla:

          Warning: you have already guessed that word
            
    * el jugador humano propone una palabra que no es coherente con las pistas dadas a propuestas previas de este mismo partido, en cuyo caso el programa imprime el siguiente mensaje en pantalla:
    
          Warning: the word you guessed is not compatible with the previous clues
    
  En los dos casos, el player hace uso de métodos de otras clases como veremos más adelante.
    
- El método `__str__` devuelve la cadena 'human'. Se llama implícitamente desde los mensajes que se imprimen en pantalla al final de la ejecución (ver el método `play` de la clase `Game`)
    

## Clases de tratamiento de datos

Surge la necesidad de tener dos clases adicionales que actúan principalmente en el tratamiento / almacenamiento de datos. 

## Métodos de las clases de tratamiento de datos

**Clase `Storage`**

La clase `Storage` representa el almacén donde se guardan las pistas vistas hasta ahora en el juego en curso.

- El método `__init__` es relativamente trivial.
- El método `add_clue` añade un objeto de la clase `Clue` al almacén y no devuelve ningún valor (por tanto, devuelve el valor especial `None`).
- El método `is_empty` devuelve `True` si el almacén no contiene ningún objeto de la clase `Clue` y `False`, en caso contrario.
- El método `is_compatible_wth_clues` devuelve `True` si al tratar la palabra pasada como argumento como palabra secreta y llamar al método `generate_clues` de la clase `Clue` para calcular la distancia entre la palabra pasada como argumento y cada candidata propuesta hasta ahora, la distancia es la misma para cada una de las palabras y pistas que la calculada con la verdadera palabra secreta. Devuelve `False` si no es el caso.

**Clase `Clue`**

La clase Clue representa la pareja (palabra candidata, pista asociada a esta palabra candidata).

- El método `__init__` es relativamente trivial.
- El *método estático* `generate_clue` devuelve un objeto de la clase `Clue` que indica la distancia entre la palabra pasada como segundo argumento, la candidata, y la palabra pasada como primer argumento, la palabra secreta. En esta práctica se tiene que tratar como en el juego *Wordle* oficial: en el caso de una letra que está presente en la palabra secreta pero aparece más veces en la palabra candidata que en la palabra secreta, las apaririciones que no están en su sitio se deben considerar como si no aparecieran.
- Los métodos `get_word` y `get_cmpnts` simplemente devuelven el valor del attributo correspondiente (se llama a este tipo de método un 'getter' o 'accesor')
- El método `is_winner` devuelve el valor `True` si los `clue_cmpnts` del objeto `self` de la clase `Clue` contiene solo valores `LETTER_IN_POS`.
- El método `__str__` simplemente llama al método adecuado de la clase `GameIO`.

**Métodos especiales de las clases `Storage` y `Clue`**

- Añadimos un método `__contains__` a la clase `Storage`: una de las advertencias del método `next_guess` de la clase `HumanPlayer` requiere averiguar si una palabra dada está contenida en un objeto de la clase `Storage`. Si el método de la clase `Storage` que realiza esta comprobación se llama exactamente `__contains__`, se puede usar la sintaxis `in` para especificar esta comprobación, p.ej. `if word in storage:`. 
- Añadimos un método `__eq__` a la clase `Clue`: la funcionalidad del método `is_compatible_with_clues` de la clase `Storage` require comprobar si dos objetos de la clase `Clue` son iguales. Si el método de la clase `Clue` que realiza esta comprobación se llama exactamente `__eq__`, se puede usar la sintaxis `==` para especificar esta comprobación, p.ej. `if clue1 == clue2:`.

<!-- grade_id: cell-5b41dd2d44e88390, points: 0, grade: false, solution: true -->
```python
#%%writefile -a wordle_class.py

class Clue(object):
    
    LETTER_IN_POS = 0
    LETTER_IN_NO_POS = 1 
    LETTER_NOT_IN = 2
    answer_cmpnts = {LETTER_IN_POS, LETTER_IN_NO_POS, LETTER_NOT_IN}
    
    def __init__(self, guess, clue_cmpnts):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __str__(self):
        return GameIO.clue_output(self)        
    
    def __repr__(self):
        return f'Clue({self.guess}, {self.clue_cmpnts})'

    
    def __eq__(self, other):
        # YOUR CODE HERE
        raise NotImplementedError()
        

    def get_word(self):
        # YOUR CODE HERE
        raise NotImplementedError()
               
    def get_cmpnts(self):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def is_winner(self):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    @staticmethod
    def generate_clue(secret_word, guessed_word):
        # YOUR CODE HERE
        raise NotImplementedError()


```

<!-- grade_id: cell-test-clue, points: 2, grade: true, solution: false -->
```python
#%%writefile -a wordle_class.py

OK = 'OK'
def test_clue():
    secret, candidate = 'mercedes', 'mercedes'
    clue = Clue.generate_clue(secret, candidate)
    print(f'{secret}, {repr(clue)}', end='...')
    assert clue.get_word() == candidate
    assert list(clue.get_cmpnts()) == [0, 0, 0, 0, 0, 0, 0, 0]
    print(OK)
    
    secret, candidate = 'mercedes', 'mercader'
    clue = Clue.generate_clue(secret, candidate)
    print(f'{secret}, {repr(clue)}', end='...')
    assert clue.get_word() == candidate
    assert list(clue.get_cmpnts()) == [0, 0, 0, 0, 2, 0, 0, 2]
    print(OK)

    secret, candidate = 'mero', 'moro'
    clue = Clue.generate_clue(secret, candidate)
    print(f'{secret}, {repr(clue)}', end='...')
    assert clue.get_word() == candidate
    assert list(clue.get_cmpnts()) == [0, 2, 0, 0]
    print(OK)
    
    secret, candidate = 'avion', 'ababa'
    clue = Clue.generate_clue(secret, candidate)
    print(f'{secret}, {repr(clue)}', end='...')
    assert clue.get_word() == candidate
    assert list(clue.get_cmpnts()) == [0, 2, 2, 2, 2]
    print(OK)

    secret, candidate = 'mero', 'mear'
    clue = Clue.generate_clue(secret, candidate)
    print(f'{secret}, {repr(clue)}', end='...')
    assert clue.get_word() == candidate
    assert list(clue.get_cmpnts()) == [0, 0, 2, 1]
    print(OK)
    
    assert Clue.generate_clue('babel', 'cauno') == Clue.generate_clue('babel', 'cauno')
    assert Clue.generate_clue('babel', 'cauno').get_cmpnts() == Clue.generate_clue('babel', 'cauto').get_cmpnts()
    assert Clue.generate_clue('babel', 'cauno').get_cmpnts() == Clue.generate_clue('babel', 'cauta').get_cmpnts()
    assert Clue.generate_clue('babel', 'cauno').get_cmpnts() == Clue.generate_clue('babel', 'cauta').get_cmpnts()
    assert not Clue.generate_clue('babel', 'cauno').is_winner()
    assert Clue.generate_clue('babel', 'babel').is_winner()

    print(OK)
    
test_clue()
```

<!-- grade_id: cell-291b9add0769b957, points: 0, grade: false, solution: true -->
```python
#%%writefile -a wordle_class.py

class Storage(object):

    def __init__(self):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def add_clue(self, clue):
        # YOUR CODE HERE
        raise NotImplementedError()

    def is_empty(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def __contains__(self, word):  #https://docs.python.org/3/reference/datamodel.html#object.__contains__
        # YOUR CODE HERE
        raise NotImplementedError()
            
    def is_compatible_with_clues(self, word):
        # YOUR CODE HERE
        raise NotImplementedError()
       
    def __repr__(self):
        # YOUR CODE HERE
        raise NotImplementedError()

```

```python

```

<!-- grade_id: cell-test-storage, points: 1, grade: true, solution: false -->
```python
#%%writefile -a wordle_class.py

def test_storage():
    storage = Storage()
    assert storage.is_empty()
    
    secret = 'folga'
    candidate = 'vinar'
    assert storage.is_compatible_with_clues(candidate)
    clue = Clue.generate_clue(secret, candidate)
    storage.add_clue(clue)
    print(storage)
    assert not storage.is_empty()
    assert storage.is_compatible_with_clues('babel')
    assert not storage.is_compatible_with_clues('clavo')
    
    candidate = 'zamba'
    assert storage.is_compatible_with_clues(candidate)
    clue = Clue.generate_clue(secret, candidate)
    storage.add_clue(clue)
    print(storage)
    assert not storage.is_compatible_with_clues(candidate)
    assert not storage.is_compatible_with_clues('clavo')
    
    candidate = 'solea'
    assert storage.is_compatible_with_clues(candidate)
    clue = Clue.generate_clue(secret, candidate)
    storage.add_clue(clue)
    print(storage)
    assert not storage.is_compatible_with_clues(candidate)
    assert not storage.is_compatible_with_clues('clavo')
    
    candidate = 'golfa'
    assert storage.is_compatible_with_clues(candidate)
    clue = Clue.generate_clue(secret, candidate)
    storage.add_clue(clue)
    print(storage)
    assert not storage.is_compatible_with_clues(candidate)
    assert not storage.is_compatible_with_clues('clavo')
    
    candidate = 'folga'
    assert storage.is_compatible_with_clues(candidate)
    clue = Clue.generate_clue(secret, candidate)
    storage.add_clue(clue)
    print(storage)
    assert not storage.is_compatible_with_clues('clavo')
    
    
    print(OK)
test_storage()
    
    
```

### Llamadas a las clases de tratamiento de datos desde las clases básicas

No sé si hace falta dar más pistas, por ejemplo:

- explicar como `get_clue` de `HumanServer` llama a `clue_input` de `GameIO`
- exlicar como `next_guess` de `ComputePlayer` llama a `is_compatible_with_clues` de `Storage`
- ...

## La clase GameIO

La clase GameIO maneja la comunicación con el mundo exterior, es decir, la entrada y la salida de datos. Se trata de una clase cuyos métodos son todos estáticos. Se suele llamar a este tipo de clase *clase de utilidad*. A continuación proporcionamos todo el código de esta clase.

```python
#%%writefile -a wordle_class.py

class GameIO:
    
    MARGIN = "          "
    L_BORDER = MARGIN + "|| "
    R_BORDER = " ||\n"
    SEPARATOR = " | "
    GO_RED = '\033[31m'
    GO_GREEN = '\033[32m'
    RESET_COL = '\033[0m'

    def __init__(self, word_length):
        pass

    @staticmethod
    def clue_cmpnts_output(clue):
        out = []
        for i, letter in enumerate(clue.get_word()):
            if clue.get_cmpnts()[i] == Clue.LETTER_IN_POS:
                out.append(GameIO.GO_GREEN + letter.upper() + GameIO.RESET_COL)
            elif clue.get_cmpnts()[i] == Clue.LETTER_NOT_IN:
                out.append(GameIO.GO_RED + letter.upper() + GameIO.RESET_COL)
            elif clue.get_cmpnts()[i] == Clue.LETTER_IN_NO_POS:
                out.append(letter.upper())
            else:
                out = ""
        return out

    @staticmethod
    def clue_output(clue):
        word_length = len(clue.get_word())
        h_border = GameIO.MARGIN + "==" + word_length * "====" + "=\n"
        return h_border + GameIO.L_BORDER \
               + GameIO.SEPARATOR.join(GameIO.clue_cmpnts_output(clue)) \
               + GameIO.R_BORDER + h_border        

    @staticmethod
    def clue_cmpnt_input(letter):
        letter = letter.lower()
        if letter == 'y':
            convert = Clue.LETTER_IN_POS
        elif letter == 'n':
            convert = Clue.LETTER_NOT_IN
        elif letter == 'e':
            convert = Clue.LETTER_IN_NO_POS     
        else:
            raise ValueError()
        return convert
       
    @staticmethod
    def clue_input(guess):
        valid_answer = False
        while not valid_answer:
            clue_string = GameIO.input(f"My guess is {guess}, what is your answer\n"\
               + "('y' for 'yes', 'n' for 'no' and 'e' for 'elsewhere', "\
               + "separated by spaces)?\n  ")
            try:
                cmpnts = tuple(map(GameIO.clue_cmpnt_input, clue_string.split()))
                if len(cmpnts)==len(guess):
                    pos = 0
                    while pos<len(cmpnts) and (cmpnts[pos] in Clue.answer_cmpnts):
                        pos += 1
                    valid_answer = pos==len(cmpnts)
            except ValueError:
                pass
            if not valid_answer:
                GameIO.output("Your answer is too long or contains invalid characters")
        return Clue(guess, cmpnts)
    
    @staticmethod
    def word_input():
        return GameIO.input('Please enter your next guess: ').lower()
    
    @staticmethod
    def secret_msg_output(word_size):
        GameIO.input(f"Press enter when you've chosen a {word_size}-letter" \
              + "word from the word base;\nbut don't tell me what it is!")
    
    @staticmethod
    def output(msg):
        print(msg)
        
    @staticmethod 
    def input(msg):
        return input(msg)
```

Observa que la entrada de las pistas proporcionadas por el usuario se hace mediante las letras 'y', 'n' y 'e', y la salida de las pistas se hace mediante letras en mayúsculas de tres colores: rojo para una letra que no esté presente en la palabra, verde para una letra que esté en el sitio correcto en la palabra y negro (normal) para una letra que esté presente en la palabra pero no esté en la posición correcta.

Finalmente, proporcionamos el código del *main* de este programa:

```python
#%%writefile -a wordle_class.py

def main():
    from os import path
    found = False
    while not found:
        dict_filename = GameIO.input(\
"""Enter the name of the dictionary file
Press enter to use the default file, i.e. lemario.txt: """\
                             )
        if not dict_filename: 
            dict_filename = "lemario.txt"
        if path.exists(dict_filename): 
            found = True
        else: 
            GameIO.output("File not found. Please try again")

    replies = {"c", "h"}
    foundp = founds = False  
    while not (foundp and founds):
        while not foundp:
            p_type = GameIO.input("Enter the player type, h for human or c for computer:\n"\
                  + "Press enter to use default player, i.e. human: ")
            if not p_type: 
                p_type = 'h'
            if p_type in replies: 
                foundp = True
            else: 
                GameIO.output("Invalid player type. Please try again.")
        while not founds:
            s_type = GameIO.input("Enter the server type, h for human or c for computer:\n"\
                  + "Press enter to use default server, i.e. computer: ")
            if not s_type: 
                s_type = 'c'
            if s_type in replies: 
                founds = True   
            else: 
                GameIO.output("Invalid server type. Please try again.")
        if p_type == s_type == "h":
            GameIO.output("Human player and human server not an option. Please try again.")
            founds = foundp = False
    
    word_length = "not_a_number"
    while word_length and not word_length.isdigit():
        word_length = GameIO.input("Enter the word length (a positive unsigned integer):\n"\
             + "Press enter to use default word length, i.e. 5: ")
    if not word_length: word_length = 5
    
    attempts = "not_a_number"
    while attempts and not attempts.isdigit():
        attempts = GameIO.input("Enter the allowed number of attempts (a positive unsigned integer):\n"\
             + "Press enter to use default word length, i.e. 6: ")
    if not attempts: attempts = 6

    a_seed = "not_a_number"
    while a_seed and not a_seed.isdigit():
        a_seed = GameIO.input("Enter the seed value (a positive unsigned integer):\n"\
             + "Press enter to let the system randomly choose a seed: ")
        # apparently, random is a global variable
    if not a_seed: 
        random.seed()
    else: 
        random.seed(a_seed)

    game = Game(dict_filename, p_type=="h", s_type=="c", word_length, attempts)
    game.play()

```

# Generación de programa

Debes generar un fichero llamado `wordle_classes.py` con el contenido de las clases presentadas y la función `main`

<!-- grade_id: cell-test-computer-player, points: 2, grade: true, solution: false -->
```python
#%%writefile -a wordle_class.py

def test_computer_player():
    wordlist = list(map(lambda x: x.strip(),
                       filter(lambda x: len(x)==6, open('lemario.txt', encoding='utf-8'))))
    random.seed(8)
    storage = Storage()
    cplay = ComputerPlayer(wordlist, storage)
    
    secret = 'folga'
    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    storage = Storage()
    cplay = ComputerPlayer(wordlist, storage)
    secret = 'cinco'
    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)


    storage = Storage()
    cplay = ComputerPlayer(wordlist, storage)
    secret = 'XXXXX'
    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    new = cplay.next_guess()
    print(f'{new}....', end='')
    assert storage.is_compatible_with_clues(new)
    print(OK)
    clue = Clue.generate_clue(secret, new)
    storage.add_clue(clue)

    try:
        new = cplay.next_guess()
        assert False
    except AssertionError as e:
        raise(e)
    except Exception as e:
        print('OK, excepción capturada')
    
    print('OK')

test_computer_player()
```

<!-- grade_id: cell-test-human-player, points: 2, grade: true, solution: false -->
```python
#%%writefile -a wordle_class.py

import io

class MyInput():
    def __init__(self, answers):
        self.pos = 0
        self.answers = answers
    def input(self, msg):
        if self.pos<len(self.answers):
            ans = self.answers[self.pos]
            self.pos += 1
        else:
            ans = ''
        return ans

def test_human_player():
    wordlist = list(map(lambda x: x.strip(),
                    filter(lambda x: len(x)==6, open('lemario.txt', encoding='utf8'))))
    old_input = GameIO.input
    old_output = GameIO.output    
    secret = 'folga'
    
    test = ['flaon', 'folia', 'folla', 'folga']
    random.seed(8)
    storage = Storage()
    hplay = HumanPlayer(wordlist, storage)
    myinput = MyInput(test)
    GameIO.input = myinput.input
    sout = io.StringIO("")
    GameIO.output = lambda x: sout.write(f'{x}')
    try:
        
        for _ in range(len(test)):
            word = hplay.next_guess()
            storage.add_clue(Clue.generate_clue(secret, word))
        assert sout.getvalue() == ''
    finally:
        GameIO.input = old_input
        GameIO.output = old_output
    
    
    test = ['flaon', 'flaon']
    random.seed(8)
    storage = Storage()
    hplay = HumanPlayer(wordlist, storage)
    myinput = MyInput(test)
    GameIO.input = myinput.input
    sout = io.StringIO("")
    GameIO.output = lambda x: sout.write(f'#{x}#')
    try:
        secret = 'folga'
        for _ in range(len(test)):
            word = hplay.next_guess()
            storage.add_clue(Clue.generate_clue(secret, word))
        #print(f':{sout.getvalue()}:')
        assert sout.getvalue() == '#Warning: you have already guessed that word#'
    finally:
        GameIO.input = old_input
        GameIO.output = old_output
    
    test = ['flaon', 'flaco']
    random.seed(8)
    storage = Storage()
    hplay = HumanPlayer(wordlist, storage)
    myinput = MyInput(test)
    GameIO.input = myinput.input
    sout = io.StringIO("")
    GameIO.output = lambda x: sout.write(f'#{x}#')
    try:
        secret = 'folga'
        for _ in range(len(test)):
            word = hplay.next_guess()
            storage.add_clue(Clue.generate_clue(secret, word))
        #print(f':{sout.getvalue()}:')
        assert sout.getvalue() == '#Warning: the word you guessed is not compatible with the previous clues#'
    finally:
        GameIO.input = old_input
        GameIO.output = old_output
        
    test = ['flaon', 'flaco', 'flaon', 'folia', 'folla', 'flaco', 'folga']
    random.seed(8)
    storage = Storage()
    hplay = HumanPlayer(wordlist, storage)
    myinput = MyInput(test)
    GameIO.input = myinput.input
    sout = io.StringIO("")
    GameIO.output = lambda x: sout.write(f'#{x}#')
    try:
        secret = 'folga'
        for _ in range(len(test)):
            word = hplay.next_guess()
            storage.add_clue(Clue.generate_clue(secret, word))
        #print(f':{sout.getvalue()}:')
        assert sout.getvalue() == '#Warning: the word you guessed is not compatible with the previous clues#'+\
            '#Warning: you have already guessed that word#'+ \
            '#Warning: you have already guessed that word#'
    finally:
        GameIO.input = old_input
        GameIO.output = old_output
        
    print('OK')
test_human_player()
    
```

<!-- grade_id: cell-test-computer-server, points: 1, grade: true, solution: false -->
```python
#%%writefile -a wordle_class.py
import io

def test_computer_server():
    wordlst = ['avion']


    old_input = GameIO.input
    old_output = GameIO.output
    try:
        cs = ComputerServer(wordlst)
        cs.generate_secret()
        assert cs.get_secret() == 'avion'
        sout = io.StringIO("")
        GameIO.output = lambda x: None
        clue = cs.get_clue('altar')
        assert clue.get_word() == 'altar' and list(clue.get_cmpnts()) == [0,2,2,2,2]

    finally:
        GameIO.input = old_input
        GameIO.output = old_output

    print('OK')

test_computer_server()

```

<!-- grade_id: cell-test-human-server, points: 1, grade: true, solution: false -->
```python
#%%writefile -a wordle_class.py

def test_human_server():
    wordlst = list(map(lambda x: x.strip(),
                       filter(lambda x: len(x)==5, open('lemario.txt', encoding='utf8'))))


    old_input = GameIO.input
    old_output = GameIO.output
    try:
        hs = HumanServer(wordlst)
        test = ['y y y b b', 'y y n n e y', 'y y y', 'y y n n e']
        myinput = MyInput(test)
        GameIO.input = myinput.input
        sout = io.StringIO("")
        GameIO.output = lambda x: sout.write(f'#{x}#')
        res = hs.get_clue('avion')
        assert res.get_word() == 'avion' and list(res.get_cmpnts()) == [0, 0, 2, 2, 1]
        assert sout.getvalue() == '#Your answer is too long or contains invalid characters##Your answer is too long or contains invalid characters##Your answer is too long or contains invalid characters##          =======================\n          || \x1b[32mA\x1b[0m | \x1b[32mV\x1b[0m | \x1b[31mI\x1b[0m | \x1b[31mO\x1b[0m | N ||\n          =======================\n#'
#         assert sout.getvalue() == \
#             """#Press enter when you've chosen a 4-letterword from the word base;
# but don't tell me what it is!#"""
    finally:
        GameIO.input = old_input
        GameIO.output = old_output

    print('OK')

test_human_server()
    
```

```python
old_random = random.random
old_randint = random.randint
```

<!-- grade_id: cell-test_play, points: 1, grade: true, solution: false -->
```python
def test_play():
    old_output = GameIO.output
    sout = io.StringIO("")
    random.seed(8)
    GameIO.output = lambda x: sout.write(f'#{x}#\n')
    try:
        game = Game('lemario.txt', False, True, 5, 100)
        game.play()
        #print(repr(sout.getvalue()))
        assert sout.getvalue() == '#          =======================\n          || \x1b[31mM\x1b[0m | O | N | \x1b[31mS\x1b[0m | \x1b[31mE\x1b[0m ||\n          =======================\n#\n#          =======================\n          || \x1b[31mI\x1b[0m | \x1b[31mN\x1b[0m | \x1b[31mG\x1b[0m | \x1b[32mO\x1b[0m | \x1b[32mN\x1b[0m ||\n          =======================\n#\n#          =======================\n          || \x1b[31mC\x1b[0m | A | L | \x1b[32mO\x1b[0m | \x1b[32mN\x1b[0m ||\n          =======================\n#\n#          =======================\n          || \x1b[32mF\x1b[0m | \x1b[32mL\x1b[0m | \x1b[32mA\x1b[0m | \x1b[32mO\x1b[0m | \x1b[32mN\x1b[0m ||\n          =======================\n#\n#The secret word was flaon.#\n#Congratulations to the computer for finding the secret word.#\n'
    finally:
        GameIO.output = old_output
    print(OK)
test_play()
```

<!-- grade_id: cell-test_program, points: 2, grade: true, solution: false -->
```python
del main
del Game
del ComputerServer
del HumanServer
del ComputerPlayer
del HumanPlayer
del Clue
del Storage
del GameIO

from wordle_classes import main, Game, ComputerPlayer, ComputerServer, HumanPlayer, HumanServer, Clue, Storage, GameIO

def test_main():
    
    test_storage()
    test_clue()
    test_computer_player()
    test_computer_server()
    test_human_player()
    test_human_server()
    test_play()
    
test_main()
```