Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

### Raices de polinomios

En una práctica anterior calculamos los ceros reales de una ecuación de tercer grado, 
en esta nos plantearemos calcular las raices de un polinomio de cualquier grado. 

Para ello vamos es interesante observar lo siguiente:

* Si el polinomio tiene grado 1 se trata de hacer una simple división
* Si $p$ es el polinomio del que queremos calcular las raices, llamamos $df$ a su derivada. Sean $x_1,x_2,\dots,x_n$ las raices ordenadas de $df$. Un estudio parecido al del caso de la ecuación de tercer grado, permite razonar sobre los posibles ceros de $p$ en $(-\infty,x_1)$, $(x_i,x_{i+1})$ y $(x_k,\infty)$ conociendo los valores: $lim_{x\rightarrow - \infty} p(x),p(x_1),p(x_2),\dots,p(x_n),lim_{x\rightarrow  \infty} p(x)$
* Los problemas surgen cuando uno de los ceros de $p$ se confunde con el de su derivada (bien porque sean iguales o porque sean muy próximos). Por ello nos vamos a empezar suponiendo que ni el polinomio $p$ ni sis derivadas tengan raices comunes. En este caso calculamos las raices de la derivada. Puesto que la derivada tiene un grado menor que el polinomio se puede proceder de forma inductiva. El caso base será cuando el grado del polinomio es 1. Será una lista como la que se indica arriba, las posibles raices de $p$ caen dentro de los intervalos mencionados.
* Para calcular lar raices dentro de los límites mencionados. Hay que comprobar que el valor del polinomio en los extremos tiene signos opuestos. Luego se procede reduciendo el intervalo a la mitad, nos quedamos con la mitad cuya evaluación en los extremos tenga signo opuesto. Procedemos de esta forma hasta que la longitud del intervalo sea más pequeño que el error $\epsilon$ considerado. En el caso de los intervalos $(-\infty, a)$ y $(a, \infty)$, lo único que necesitamos es saber el signo del $lim_{x\rightarrow - \infty} p(x)$ y $lim_{x\rightarrow \infty} p(x)$. El primero viene determinado por el grado del polinomio y el signo de coeficiente de mayor grado, el segundo solo por este último.

Normalmente llamamos "libre de cuadrados" a un polinomio sin raices múltiples.

En esta práctica usaremos polinomios con coeficientes racional. Por ello una solución $s$ será un intervalo un par de racionales $li$, $ls$ tales que la solución cae dentro del intervalo $[li, ls]$ y $ls-li<\epsilon$ para un $\epsilon$ dado.

Hay ficheros que implementan los números racionales `rationals.py`, monomios `monomials.py`. Consulta las implementaciones e intenta comprender como funcionan. A partir de ellos vamos a construir la clase `Polynomials`. A lo mejor es interesante que realizes y compruebes tu implementación en un IDE externo (como spyder) y luego la incluyas aquí. La entrega consiste en este cuaderno completo más el fichero con la implementación de los polinomios que realices que debe estar en un fichero llamado `polinomials.py`.

Debes implementar las funciones que faltan:
- Suma, multiplicación y división de polinomios: métodos `__add__`, `__mul__`, y `__divmod_`. Consulta la [documentación de python](https://docs.python.org/3/reference/datamodel.html#emulating-numeric-types) para ver cómo deben funcionar estos métodos
- Función `gcd` que calcula el máximo común divisor de dos números. Se debe implementar de forma recursiva. Puesto que los polinomios tienen la comparación con el cero y la división con resto, esta misma función debe funcionar con los polinomios cuando se completen las operaciones aritméticas.
- Derivada de un polinomio: método `diff`
- Evaluar el polinomio en un punto: método `eval`
- Una método que devuelve los signo de $lim_{x\rightarrow - \infty} p(x)$ y $lim_{x\rightarrow \infty} p(x)$: `sgn_lim_minus_inf` `sgn_lim_inf`. Devuelven -1 si el signo es negativo y 1 en caso contrairo
- Sabiendo que el valor del polinomio $p$ en dos puntos tienen signos opuestos, hacer una función que calcule una solución del polinomio entre esos dos puntos con un error $\epsilon$. Una solución sera un par de racionales $a,b$ con $a\leq b$ de forma que $b-a<\epsilon$ y la solución real está en ese intervalo. La forma de encontrar la solución es parecida a la de la búsqueda dicotómica. Se parte de el intervalo original y se calcula el punto medio. La solución estará en uno de los 2 subintervalos. Método: zero_bolzano. El método debe lanzar una excepción si el valor del polinomio en los dos puntos tienen el mismo signo.
- Si tenemos una cota inferior de la solución pero desconocemos la superior, ésta la podemos calcular de la siguiente forma. Tomamos un número positivo superior a la cota conocida. Si el signo del polinomio en el candidato es distino al valor del polinomio en la cota, lo tomamos como cota superior (estamos en las hipótesis del teorema de Bolzano). Si no lo es, duplicamos el valor del candidado hasta que éste cumpla la condición anterior. Ahora podemos aplicar el método anterior para encontrar la soluición: método `zero_inf`. El método debe lanzar una excepción en caso que el signo del polinomio en el punto coincida con el signo de $lim_{x\rightarrow \infty} p(x)$.
- De forma parecida se puede proceder si se conoce una cota superior pero no inferior. Método `zero_minusinf`. 
El método debe lanzar una excepción en caso que el signo del polinomio en el punto coincida con el signo de $lim_{x\rightarrow -\infty} p(x)$.

- Ahora podemos aplicar Método descrito arriba para calcular las raices de un polinomio sin raices múltiples. Método `zeroes_nomultiple`.
- Para calcular las raices de un polinomio genérico calculamos el máximo común divisor entre el polinomio y su derivada. Dividimos nuestro polinomio entre el máximo común divisor y calculamos las raíces del nuevo polinomio calculado. Método `zeroes`.

<!-- grade_id: cell-bbcb2e48e92e1ebe, points: 0, grade: false, solution: true -->
```python
from rationals import Rational
import rationals
from monomials import Monomial
from copy import copy, deepcopy #See http://docs.python.org/library/copy.html
import sys
import random


def gcd(a,b):
    """
    The Eucludes algorithm for The greatest common divisor, 
    it should run on polynomials. 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
def simplify(monomials):
    """
    makes a simplified version of the monomials
    - makes an independent copy. The copy does not share memory with the original list
    - the monomials are sorted
    - there are no two monomials with the same degree
    - there are no monomials the zoer coeficient
    """
    monomials = deepcopy(monomials)
    monomials.sort()
    return remove_zero_monomials(simplify_equal_degree_monomials(monomials))


def remove_zero_monomials(monomials):
    """
    This function remove all monomials whose coefficient is zero
    """
    return list(filter(lambda x: x.coef != rationals.ZERO, monomials))


def simplify_equal_degree_monomials(monomials):
    """
    monomials is list of monomials sorted according the
    degree.
    """
    def simplify_equal_degree_monomials_rec(monomials, pos):
        if pos < 0:
            return []
        else:
            mlst = simplify_equal_degree_monomials_rec(monomials, pos-1)
            mon = monomials[pos]
            if mlst == [] or mlst[-1].degree!=mon.degree:
                mlst.append(mon)
            else:
                mlst[-1] += mon
            return mlst

    return simplify_equal_degree_monomials_rec(monomials, len(monomials)-1)


def test_remove_zero_monomials():
    tests = [
        ([], []),
        (["2·X^2", "0·X^3", "1·X^1", "3"], ["2·X^2", "1·X^1", "3"]),
        (["0·X^2", "3·X^3", "1·X^1", "3"], ["3·X^3", "1·X^1", "3"]),
        (["2·X^2", "3·X^3", "1·X^1", "0"], ["2·X^2", "3·X^3", "1·X^1"])
    ]

    for t, s in tests:
        t = list(map(Monomial.parse, t))
        s = list(map(Monomial.parse, s))
        res = remove_zero_monomials(t)
        assert res==s, f'Fail {t} {s} {res}'
    print('OK. remove_zero_monomials')


def test_simplify_equal_degree_monomials():
    import random
    tests = [
        ([], []),
        (["3·X^3", "2·X^2", "1·X^1", "7", "3"], ["3·X^3", "2·X^2", "1·X^1", "10"]),
        (["3·X^3", "2·X^2", "1·X^1", "5", "7", "3"], ["3·X^3", "2·X^2", "1·X^1", "15"]),
        (["3·X^3", "2·X^2", "8·X^1", "6·X^1", "4·X^1", "2·X^1", "5", "7", "3"], ["3·X^3", "2·X^2", "20·X^1", "15"]),
        (["9·X^3", "6·X^3", "3·X^3", "2·X^2", "8·X^1", "6·X^1", "4·X^1", "2·X^1", "5", "7", "3"], ["18·X^3", "2·X^2", "20·X^1", "15"]),
    ]

    for t, s in tests:
        t = list(map(Monomial.parse, t))
        s = list(map(Monomial.parse, s))
        res = simplify_equal_degree_monomials(t)
        assert res==s, f'Fail {t} {s} {res}'
    print('OK. simplify_equal_degree_monomials')


def test_simplify():
    tests = [
        ([], []),
        (["3·X^3", "2·X^2", "1·X^1", "7", "3"], ["3·X^3", "2·X^2", "1·X^1", "10"]),
        (["3·X^3", "2·X^2", "1·X^1", "5", "7", "3"], ["3·X^3", "2·X^2", "1·X^1", "15"]),
        (["3·X^3", "2·X^2", "8·X^1", "6·X^1", "4·X^1", "2·X^1", "5", "7", "3"], ["3·X^3", "2·X^2", "20·X^1", "15"]),
        (["9·X^3", "6·X^3", "3·X^3", "2·X^2", "8·X^1", "6·X^1", "4·X^1", "2·X^1", "5", "7", "3"], ["18·X^3", "2·X^2", "20·X^1", "15"]),
        (["9·X^3", "6·X^3", "3·X^3", "-6·X^3", "-3·X^3", "-6·X^3", "-3·X^3", "2·X^2", "8·X^1", "6·X^1", "4·X^1", "2·X^1", "5", "7", "3"], ["2·X^2", "20·X^1", "15"]),
        (["9·X^3", "6·X^3", "3·X^3", "-6·X^3", "-3·X^3", "-6·X^3", "-3·X^3", "2·X^2", "8·X^1", "6·X^1", "4·X^1", "2·X^1", "5", "7", "3", "-2·X^2", "-20·X^1", "-15"], []),
    ]

    for t, s in tests:
        random.shuffle(t)
        t = list(map(Monomial.parse, t))
        s = list(map(Monomial.parse, s))
        s.sort()
        res = simplify(t)
        assert res==s, f'Fail {t} {s} {res}'
    print('OK. simplify')


class Polynomial(object):
    """
    This class represent a polynomial with rational coefficients.
    The polynomial is represented by a list of monomials satisfying:
    * There are no monomials with zero coeficient.
    The zero polynomial is the empty list
    * The list is ordered according the degree of the monomial.
    The longer degree first.
    * There are no monomials with the same degree.
    """
    def __init__(self, monomials=[]):
        # If we do not copy the monomials and we modify the monomials
        # the default parameter may change see
        self.monomials = simplify(monomials)
        
    @staticmethod
    def parse(s):
        """
        This static function parses a polynomial from the string s.
        The string maybe the empty one (to denote the zero polynomial)
        of have the following form
        <monomial> + <monomial> + <monomial> .... + <monomial>
        @type s: string
        @rtype: Polynomial
        @raise Exception: if the monomial is not as described above

        """
        l = s.strip().split("+")
        monomials = [None]*len(l)
        i=0
        while i<len(l):
            monomials[i] = Monomial.parse(l[i])
            i+=1
        return Polynomial(monomials)


    def normalize(self):
        """
        Returns an new polinomial normalized with the most representative
        coeficient equals to one.
        """
        monomials = []
        if len(self.monomials)>0:
            coef = self.monomials[0].coef
            i = 0
            while i<len(self.monomials):
                m = deepcopy(self.monomials[i])
                m.coef = m.coef/coef
                monomials.append(m)
                i += 1
        return Polynomial(monomials)


    def __mul__(self,other):
        # YOUR CODE HERE
        raise NotImplementedError()

    def __add__(self,other):
        # YOUR CODE HERE
        raise NotImplementedError()

    def __neg__(self):
        """
        This function returns the opposite polynomial
        @rtype: Polynomial
        """
        n = deepcopy(self)
        for m in n.monomials:
            m.coef = -m.coef
        return n

    def __sub__(self,other):
        """
        This function implements the subtraction of polynomials
        @type other: polynomial
        @rtype: Polynomial
        """
        return self + (-other)

    def max_degree(self):
        """
        This function returns the coeficient of the monomial of degree 0 of the polynomial
        @rtype: Rational
        """
        n=len(self.monomials)
        if n==0:
            return Monomial(Rational(0),0)
        else:
            return self.monomials[n-1]

    def degree(self):
        """
        This fucntion returns the degree of the polynomial
        @rtype: integer
        """
        return self.max_degree().degree

    def div_zero_degree(self,q):
        """
        This function divides polynomial p by polynomial q, but
        having into acount that q has degree zero. The remainder is
        always zero, so this function only returns a polynomial
        @type q: Polynomial
        @param q: zero degree polynomial
        @rtype Polynomial
        @param: d where self = q * d
        """
        if len(q.monomials)==0:
            raise Exception("The divisor cannot be zero")
        else:#if degree of q is zero and is not the polynomial zero
            #the only monomial it can have has the coieficient that has
            #to divide q and it cannot be zero
            r = q.monomials[0].coef
        d = deepcopy(self)
        for m in d.monomials:
            m.coef = m.coef/r
        return d

    def __divmod__(self, other):
        """
        This function implements the divmod function
        @rtype other: polynomial
        @rtype: tuple of two polynomials
        @param: returns (q,r) such that
          - if other.degree() >0: self= q * other + r and r.degree()<other.degree()
          - if other.degree()==0: seff = q*other and r is the zero polynomial
        @raise Exception: if the degree of other is zero
        """
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def __floordiv__(self,other):
        """
        This function implements the // operation
        @rtype other: polynomial
        @rtype: polynomial
        @param: returns the quiotient of the division see __divmod()__
        @raise Exception: if the degree of other is zero
        """
        (q,r) = divmod(self,other)
        return q

    def __mod__(self,other):
        """
        This function implements the % operation
        @rtype other: polynomial
        @rtype: tuple of two polynomials
        @param: returns the remainder of the division see __divmod()__
        @raise Exception: if the degree of other is zero
        """
        (q,r) = divmod(self,other)
        return r

    def __repr__(self):
        if len(self.monomials)>0:
            s = ""
            for m in self.monomials:
                if s!="":
                    s+=" + "
                s += str(m)
            return s
        else:
            return "0"


    """Functions to compute equality. We need to compare the polynomial
       to the zero integer to apply the euclidean_domain functions.
       So we extend the equality to any numerical value,
       so the rationals must implement the equality to any numerical value"""
    def eq_num(self,num):
        if len(self.monomials) > 1:
            return False
        elif len(self.monomials) == 1:
            m = self.monomials[0]
            return m.degree == 0 and m.coef == num
        else:
            return num == 0

    def eq_poly(self, other):
        n = len(self.monomials)
        m = len(other.monomials)
        if m != n:
            return False
        i=0
        while i<n and self.monomials[i]==other.monomials[i]:
            i+=1
        return i==n

    def __eq__(self,other):
        if isinstance(other,Polynomial):
            return self.eq_poly(other)
        elif isinstance(other,int):
            return self.eq_num(other)
        else:
            return NotImplemented

    def diff(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def eval(self, x):
        # YOUR CODE HERE
        raise NotImplementedError()

    def sgn_lim_inf(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def sgn_lim_minus_inf(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def zero_inf(self, lb, epsilon):
        # YOUR CODE HERE
        raise NotImplementedError()


    def zero_minusinf(self, ub, epsilon):
        # YOUR CODE HERE
        raise NotImplementedError()


    def zero_bolzano(self, lb, ub, epsilon):
        """
        Finds a solution of the polinomial in the interval
        [li,ub]. The polynomial, li, and ub meet the conditions of
        Bolzano's Theorem: self.eval(li)*self.eval(ub)<0

        A solution is a rational interval [a,b] such that b-a<epsilon
        """
        # YOUR CODE HERE
        raise NotImplementedError()




    def zeroes_nomultiple(self, epsilon):
        """
        Neither the polinomial or its derivatives
        have roots with multiplicity greatger than 1 
        
        Returns an order list of solutions. Each solutions
        is of the form [li, ls] with ls-li<epsilon
        """
        # YOUR CODE HERE
        raise NotImplementedError()

    def zeroes(self, epsilon):
        # YOUR CODE HERE
        raise NotImplementedError()



def test_init():
    a = Polynomial([Monomial(coef=Rational(1)),
                    Monomial(coef=Rational(2)),
                    Monomial(coef=Rational(0), degree=1)])
    assert repr(a) == "3"

    a = Polynomial([Monomial(Rational(1),1),
                    Monomial(Rational.parse("2"),2),
                    Monomial(Rational.parse("-1/3"),1),
                    Monomial(Rational.parse("-2"),2),
                    ])
    assert repr(a) == "2/3·x^1"
    print("OK")










def test_eq():
    a = Polynomial([Monomial(coef=Rational(1)),
                    Monomial(coef=Rational(2)),
                    Monomial(coef=Rational(0), degree=1)])
    b = Polynomial.parse("3")
    assert a == b

    a = Polynomial([Monomial(Rational(1),1),
                    Monomial(Rational.parse("2"),2),
                    Monomial(Rational.parse("-1/3"),1),
                    Monomial(Rational.parse("-2"),2),
                    ])
    assert a == Polynomial.parse("2/3·x^1")

    a = Polynomial([Monomial(Rational(1),1),
                    Monomial(Rational.parse("2"),2),
                    Monomial(Rational.parse("-1/3"),1)
                    ])
    assert a == Polynomial.parse("2·x^2 + 2/3·x^1")

    a = Polynomial([Monomial(Rational(1),1),
                    Monomial(Rational.parse("2"),2),
                    Monomial(Rational.parse("-1/3"),1)
                    ])
    assert a == Polynomial.parse("2·x^2 + 2/3·x^1 + 0·x^5")

    assert a != Polynomial.parse("2·x^2 + 1/3·x^1 + 0·x^5")
    assert a != Polynomial.parse("2·x^2 + 2/3·x^1 + 2·x^5")
    assert a != Polynomial.parse("1·x^2 + 2/3·x^1 + 2·x^5")

    print('OK')







```

```python
test_eq()
```

<!-- grade_id: cell-1f61165ad34f25d8, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""

def test_sum():
    p1 = Polynomial.parse("2·x^3 + -2/3·x^1 + 2·x^1")
    p2 = Polynomial.parse("3·x^4 + 2/3·x^1 + -1·x^1")
    p3 = p1 + p2
    assert p3 == Polynomial.parse("3·x^4 +2·x^3 + 1·x^1")

    p1=Polynomial.parse("5 +           3·x^2 + 4·x^3")
    p2=Polynomial.parse("6 + 1·x^1 + -3·x^2 + 6·x^4")
    p3 = p1 + p2
    assert p3 == Polynomial.parse("6·x^4 +4·x^3 + 1·x^1 + 11")
    print('OK')

test_sum()
```

<!-- grade_id: cell-7314ce67b82e8631, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""


def test_mul():
    p1 = Polynomial.parse("2·x^3 + -2/3·x^1 + 2·x^1")
    p2 = Polynomial.parse("3·x^4 + 2/3·x^1 + -1·x^1")
    p3 = p1 * p2
    print(p3)
    assert p3 == Polynomial.parse("-4/9·x^2 + -2/3·x^4 + 4·x^5 + 6·x^7")

    p1=Polynomial.parse("5 +           3·x^2 + 4·x^3")
    p2=Polynomial.parse("6 + 1·x^1 + -3·x^2 + 6·x^4")
    p3 = p1 * p2
    print(p3)
    assert p3 == Polynomial.parse("30 + 5·x^1 + 3·x^2 + 27·x^3 + 25·x^4 + -12·x^5 + 18·x^6 + 24·x^7")
    print('OK')

test_mul()
```

<!-- grade_id: cell-2b81a2bbd89ea57f, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""


def test_div():
    c = Polynomial.parse("5 +           3·x^2 + 4·x^3")
    d = Polynomial.parse("6 + 1·x^1 + -3·x^2 + 6·x^4")
    r = Polynomial.parse("1/2 + -3/4·x^1 + 1/3·x^2")
    D = d*c + r
    print(D)

    d1, r1 = divmod(D, c)
    assert d1 == d and r1 == r

    print('OK')
test_div()
```

<!-- grade_id: cell-74ef7236bad10170, points: 1, grade: true, solution: false -->
```python
def test_gcd():
    a = Polynomial.parse("5 +           3·x^2 + 4·x^3")
    b = Polynomial.parse("6 + 1·x^1 + -3·x^2 + 6·x^4")
    c = gcd(a,b)
    assert c.degree() ==  0

    a = Polynomial.parse("1·x^2 + 1")
    b = Polynomial.parse("1·x^2 + -1")
    c = Polynomial.parse("1·x^1 + 2")

    d = a * c
    e = b * c

    f = gcd(d, e)

    assert c.normalize() == f.normalize()
    print('OK gcd')
test_gcd()
```

<!-- grade_id: cell-b517f1ae8b49709d, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""

def test_diff():
    a = Polynomial.parse("6 + 1·x + -3·x^2 + 6·x^4")
    b = Polynomial.parse("1 + -6·x + 24·x^3")
    c = Polynomial.parse("-6 + 72·x^2")
    d = Polynomial.parse("144·x")
    e = Polynomial.parse("0")
    assert a.diff() == b
    assert b.diff() == c
    assert c.diff() == d
    assert e.diff() == e
    assert e.diff() == e
    print('OK. diff')
    
test_diff()
```

<!-- grade_id: cell-7171d5b71f9fd583, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""


def test_eval():
    p1 = Polynomial.parse("1/2 + 2·x")
    assert p1.eval(Rational(0,1)) == Rational(1,2)
    assert p1.eval(Rational(1,1)) == Rational(5,2)
    assert p1.eval(Rational(1,2)) == Rational(3,2)

    p1 = Polynomial.parse("1/2 + 2·x + -1·x^2")
    assert p1.eval(Rational(0,1)) == Rational(1,2)
    assert p1.eval(Rational(1,1)) == Rational(3,2)
    assert p1.eval(Rational(1,2)) == Rational(5,4)

    p1 = Polynomial.parse("1/2 + 2·x + -1·x^2 + 1/3·x^3")
    assert p1.eval(Rational(0,1)) == Rational(1,2)
    assert p1.eval(Rational(1,1)) == Rational(11,6)
    assert p1.eval(Rational(1,2)) == Rational(31,24)

    print('Ok. eval')

test_eval()
```

<!-- grade_id: cell-d37bbce1771b5dff, points: 1, grade: true, solution: false -->
```python
def test_sgn_lim_inf():
    tests = [
        ("-2 + 1·x^2", 1 ),
        ("-2 + -1·x^2", -1 ),
        ("-2 + -1·x^2 + 3·x^3", 1 ),
        ("-2 + -1·x^2 + -3·x^3", -1 )
    ]
    for p, s in tests:
        p = Polynomial.parse(p)
        assert p.sgn_lim_inf() == s, f"Fail: {p} {p.sgn_lim_inf()}"

    print('ok sgn_lim_inf')
    
test_sgn_lim_inf()
```

<!-- grade_id: cell-c70d0aed6ca332d0, points: 1, grade: true, solution: false -->
```python

def test_sgn_lim_minus_inf():
    tests = [
        ("-2 + 1·x^2", 1 ),
        ("-2 + -1·x^2", -1 ),
        ("-2 + -1·x^2 + 3·x^3", -1 ),
        ("-2 + -1·x^2 + -3·x^3", 1 )
    ]
    for p, s in tests:
        p = Polynomial.parse(p)
        assert p.sgn_lim_minus_inf() == s, f"Fail: {p} {p.sgn_lim_inf()}"

    print('ok sgn_lim_minus_inf')
    
test_sgn_lim_minus_inf()
```

<!-- grade_id: cell-100cf9e75d68412c, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""

def test_bolzano():
    import math
    tests = [
        ("1/2 + 2·x", Rational(-1,1), Rational(0,1), -0.25),
        ("-2 + 1·x^2", Rational(1,1), Rational(2,1), math.sqrt(2))
    ]
    epsilon = 1.0e-5

    for p, li, ls, sol in tests:
        p = Polynomial.parse(p)
        li,ls = p.zero_bolzano(li, ls, epsilon)
        print(ls, li, li.to_float(), ls.to_float(), sol)
        assert ls-li<epsilon
        assert li<=sol<=ls

    print('ok. bolzano')
    
test_bolzano()
```

<!-- grade_id: cell-b53b36d892bc6164, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""

def test_zero_inf():
    import math
    tests = [
        ("1/2 + 2·x", Rational(-1,1), -0.25),
        ("-2 + 1·x^2", Rational(1,1), math.sqrt(2)),
        ("-2 + 1·x^2", Rational(-1,1), math.sqrt(2)),
        ("1 + 1·x", Rational(-2,1), -1)

    ]
    epsilon = 1.0e-5

    for p, li, sol in tests:
        p = Polynomial.parse(p)
        li,ls = p.zero_inf(li, epsilon)
        print(ls, li, li.to_float(), ls.to_float(), sol)
        assert ls-li<epsilon, f'Fail {p} {li} {sol}'
        assert li<=sol<=ls, f'Fail {p} {li} {sol}'

    print('ok. zero inf')

test_zero_inf()
```

<!-- grade_id: cell-ff0acf2d27261b72, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""

def test_zero_minusinf():
    import math
    tests = [
        ("1/2 + 2·x", Rational(1,1), -0.25),
        ("-2 + 1·x^2", Rational(1,1), -math.sqrt(2)),
        ("-2 + 1·x^2", Rational(-1,1), -math.sqrt(2)),
        ("1 + 1·x", Rational(-1,2), -1)
    ]
    epsilon = 1.0e-5

    for p, ls, sol in tests:
        p = Polynomial.parse(p)
        li,ls = p.zero_minusinf(ls, epsilon)
        print(ls, li, li.to_float(), ls.to_float(), sol)
    for p, ls, sol in tests:
        p = Polynomial.parse(p)
        li,ls = p.zero_minusinf(ls, epsilon)
        print(ls, li, li.to_float(), ls.to_float(), sol)
        assert ls-li<epsilon, f'Fail {p} {ls} {sol}'
        assert li<=sol<=ls, f'Fail {p} {ls} {sol}'

    print('ok. zero minus inf')
test_zero_minusinf()
```

<!-- grade_id: cell-4453c81ff1dcee09, points: 2, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
"""

def test_zeroes_nomultiple():
    import math
    tests = [
        ("1/2 + 2·x", [-0.25]),
        ("-2 + 1·x^2", [-math.sqrt(2), math.sqrt(2)]),
        ("-1 + -4·x^1 + 1/2·x^2 + 2·x^3", [-math.sqrt(2), -0.25, math.sqrt(2)]),
        ("3 + 11·x^1 + -11/2·x^2 + -11/2·x^3 + 2·x^4", [-math.sqrt(2), -0.25, math.sqrt(2), Rational(3,1)]),
    ]
    epsilon = 1.0e-5

    for p, real_sols in tests:
        p = Polynomial.parse(p)
        print(f'testing {p}.......')
        sols = p.zeroes_nomultiple(epsilon)
        assert len(sols) == len(real_sols), f"{p} should have {len(real_sols)} solutions, the function has found {len(sols)}"
        i = 0
        while i<len(sols):
            li, ls = sols[i]
            #print(ls, li, li.to_float(), ls.to_float(), real_sols[i])
            assert ls-li<epsilon, f'Fail {p} {real_sols[i]} {li} {ls} {sols[i]}'
            assert li<=real_sols[i]<=ls, f'Fail {p} {real_sols[i]} {li} {ls} {sols[i]}'
            i+=1
        print(f'{p} OK.......\n.....')

    print('ok zeros nomultiple')

test_zeroes_nomultiple()


```

<!-- grade_id: cell-b89f155c8df2957d, points: 3, grade: true, solution: false -->
```python
"""
La implementación de forma recursiva da derecho a 2 puntos.
Si se implementa de forma iterativa es 1 punto
Para obtener 3 puntos es necesario que zeroes y 
zeroes_nomultiple se llamen de forma mútua de forma adecuada. 
Es decir que zeroes no sea una copia de la zeroes_nomultiple dividiendo por la derivada
cuando es preciso.
"""

def test_zeroes():
    import math
    tests = [
        ("1/2 + 2·x", [-0.25]),
        ("-2 + 1·x^2", [-math.sqrt(2), math.sqrt(2)]),
        ("-1 + -4·x^1 + 1/2·x^2 + 2·x^3", [-math.sqrt(2), -0.25, math.sqrt(2)]),
        ("3 + 11·x^1 + -11/2·x^2 + -11/2·x^3 + 2·x^4", [-math.sqrt(2), -0.25, math.sqrt(2), Rational(3,1)]),
        ("9 + 66·x^1 + 88·x^2 + -154·x^3 + -315/4·x^4 + 209/2·x^5 + 33/4·x^6 + -22·x^7 + 4·x^8", [-math.sqrt(2), -0.25, math.sqrt(2), Rational(3,1)]),
        ("-1 + 1·x^3", [1.0]),
        ("1 + 3·x^1 + 3·x^2 + 1·x^3", [-1]),
        ("1 + 4·x^1 + 6·x^2 + 4·x^3 + 1·x^4", [-1])
    ]
    epsilon = 1.0e-5

    for p, real_sols in tests:
        p = Polynomial.parse(p)
        print(f'testing {p}.......')
        sols = p.zeroes(epsilon)
        assert len(sols) == len(real_sols)
        i = 0
        while i<len(sols):
            li, ls = sols[i]
            print(ls, li, li.to_float(), ls.to_float(), real_sols[i])
            assert ls-li<epsilon, f'Fail {p} {real_sols[i]} {li} {ls} {sols[i]}'
            assert li<=real_sols[i]<=ls, f'Fail {p} {real_sols[i]} {li} {ls} {sols[i]}'
            i+=1
        print(f'{p} OK.......\n.....')

    print('ok zeros')

test_zeroes()
```